# Homework 3: Conditioning and Stability

(Note that in this homework we use $\epsilon_m$ where the textbook uses $\epsilon_{\text{machine}}$ to save some keystrokes.)

```julia
import Pkg; Pkg.add("Plots");
using LinearAlgebra; using Plots; default(fmt=:png);
```

**1 (T&B 12.3).** Consider two classes of random matrices.  The first class has entries that are independent samples from the real normal distribution with mean zero and standard distribution $m^{-1/2}$. (Note that julia's `randn` produces matrices like this, but the entries have standard distribution $1$.)  The second class of matrices are the same as the first but upper triangular.

**(a) [4%]** For the first class of matrices for $m = {8, 16, 32, \dots, 1024}$, create 10 random matrices.  Create four scatter plots:

1. A scatter plot of all of the eigenvalues of all the matrices in the complex plane on top of each other.

2. A scatter plot of the spectral radius $\rho(A)$ of all the matrices, separated so that the trend in $m$ can be seen.

3. A scatter plot of the 2-norms of all the matrices, separated so that the trend in $m$ can be seen.

4. A scatter plot of $\sigma_m$ of all the matrices, separated so that the trend in $m$ can be seen.

In the next question you will estimate any trends that may be present, so scale your plots appropriately.

<!-- grade_id: cell-2bb8d5fc87bc0e5a, points: 4, grade: true, solution: true -->
```julia
# your code here
```

**(b) [3%]** Look at your plots from part (a) and estimate the trends of $\rho(A)$, $\|A\|_2$, and $\sigma_m(A)$. (If a value is convergening to a constant, what constant does it look like it's converging to?  If it is increasing or decreasing, is there a simple function of $m$ that it is behaving like?)

<!-- grade_id: cell-9922ce25748381d4, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

**(c) [4%]** Create four similar plots for the class of random upper triangular matrices described.  For the scatter plot of the eigenvalues, the eigenvalues will all be real, so do not plot them in the complex plane.  Instead, create a plot like your other three with the values separated by $m$ so that you can see the trend.

<!-- grade_id: cell-919459baa3e7dd44, points: 4, grade: true, solution: true -->
```julia
# your code here
```

**(d) [4%]** As in part (b), estimate the trends that you see in part (c). (If any of the values fall below $\epsilon_m$, see if you can estimate the trend from the data before that happens.)

<!-- grade_id: cell-4bbb9c11bfb83bef, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

**2 (T&B 13.3).** Consider the polynomial $p(x) = (x-2)^9 = x^9 - 18 x^8 + 144 x^7 - 672 x^6 + 2016x^5 - 4032 x^4 + 5376 x^3 - 4608 x^2 + 2304 x - 512.$

**(a) [3%]** Plot $p(x)$ on the points `x = 1.930:0.001:2.070`, evaluating $p$ via its coefficients $1$, $-18$, $\dots$.

<!-- grade_id: cell-a153711200aff5ea, points: 3, grade: true, solution: true -->
```julia
# your code here
```

**(b) [3%]** Produce the same plot again, now evaluating $p$ via the expression $(x - 2)^9$.

<!-- grade_id: cell-a897af31923a73f3, points: 3, grade: true, solution: true -->
```julia
# your code here
```

**3 (T&B 14.2).** 

**(a) [6%]** Show that $(1 + O(\epsilon_m))(1 + O(\epsilon_m)) = 1 + O(\epsilon_m).$  
The precise meaning of this statement is that if $f$ is a function satisfying
$f(\epsilon_m) = (1 + O(\epsilon_m))(1 + O(\epsilon_m))$
as $\epsilon_m \to 0$,
then $f$ also satisfies $f(\epsilon_m) = 1 + O(\epsilon_m)$ as $\epsilon_m \to 0.$

<!-- grade_id: cell-504d088674f66bc8, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) [6%]** Show that $(1 + O(\epsilon_m))^{-1} = 1 + O(\epsilon_m).$

<!-- grade_id: cell-88fe80029488a522, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

**4.** When analyzing a computation that is the result of a large numbers of floating point operations, we often end up with expressions like

$$\hat{x} = x(1 + \delta_1)(1 + \delta_2)\cdots(1 + \delta_n),$$

where $x$ is the exact value and there have been $n$ relative errors $\delta_1, \dots, \delta_n$, each with magnitude $|\delta_i| \leq \epsilon_m$. An upperbound on the error is

$$\frac{|x - \hat{x}|}{|x|} \leq \underbrace{(1 + \epsilon_m)^n - 1}_{e_n(\epsilon_m)}.$$

**(a) [6%]** We could prove that the quantity $e_n(\epsilon_m)$ is $O(\epsilon_m)$, but we can do one better.  Show that $e_n(\epsilon_m) = n \epsilon_m + O(\epsilon_m^2).$

<!-- grade_id: cell-2d2d40a6576caf6e, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) [6%]** The quantity $e_n(\epsilon_m)$ is hard to work with, and the asymptotic approximation $n\epsilon_m$ from part (a) is not an upper bound.

It turns out that an upper bound of $e_n(\epsilon_m)$ can be defined as

$$
\gamma_n(\epsilon_m) = \frac{n\epsilon_m}{1 - n\epsilon_m},$$

which is a good upper bound as long as $n \epsilon_m \ll 1.$

Prove that $\gamma_n(\epsilon_m) \geq e_n(\epsilon_m)$ whenever $n\epsilon_m < 1.$

<!-- grade_id: cell-7be293dc98049fce, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

**c [3%]** Create a plot showing $e_n$ and $\gamma_n$ for $\epsilon_m$ for 32-bit floating point numbers for $n = 0$ to 1 million.

**5 bonus points.** Determine a rational function

$$
\eta_n(\epsilon_m) = \frac{n \epsilon_m}{\alpha_n + \beta_n \epsilon_m},
$$
where $\alpha_n$ and $\beta_n$ can be any functions of $n$, such that

$$\eta_n(\epsilon_m) \geq e_n(\epsilon_m), \quad \text{ whenever }\alpha_n + \beta_n \epsilon_m < 1,$$

and such that $\eta_n(\epsilon_m) - e_n(\epsilon_m) = O(\epsilon_m^3)$, and add $\eta_n$ to your plot.

```julia
emach = Float64(eps(Float32))
```

<!-- grade_id: cell-adb48dfe564f005f, points: 3, grade: true, solution: true -->
```julia
# your code here
```

**5 (T&B 15.2).** Consider an algorithm for the problem of computing the (full) SVD of a matrix.  The data for this problem is a matrix $A$, and the solution is three matrices $U$ (unitary), $\Sigma$ (diagonal), and $V$ (unitary) such that $A = U \Sigma V^*$.  (We are speaking here of explicit matrices $U$ and $V$, not implicit representations as products of reflectors.)

**(a) [3%]** Explain what it would mean for this algorithm to be backward stable.

<!-- grade_id: cell-caed35ddc9a83aea, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) [5%]** In fact, for a simple reason, this algorithm cannot be backward stable.  Explain.

<!-- grade_id: cell-64adf8ee8c37d216, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

**(c) [5%]** Fortunately, the standard algorithms for computing the SVD are stable.  Explain what stability means for such an algorithm.

<!-- grade_id: cell-6378cca657d4f133, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

**6 (Inspired by T&B 16.1).** Let $Q$ and $A$ be $n \times n$ floating point matrices (meaning $Q = \mathrm{fl}(Q)$ and $A = \mathrm{fl}(A)$), and suppose $Q$ is unitary.  Define $f(A) = QA$ in exact arithmetic, and let $\tilde{f}(Q) = \mathrm{fl}(QA)$ computed in the straightforward way.

**(a) [7%]** Show that $\tilde{f}$ is accurate in the Frobenius norm,

$$\frac{\|\tilde{f}(A) - f(A)\|_F}{\|f(A)\|_F} = O(\epsilon_m).$$

You may use the backward stability of the inner product without proof.

<!-- grade_id: cell-7a7ca0ee9dbb54e2, points: 7, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) [7%]** Use your answer from part (a) to show that $\tilde{f}$ is backward stable.

_Hint._ Use $Q^*(\tilde{f}(A) - f(A))$.

<!-- grade_id: cell-8eebf256ea5130ec, points: 7, grade: true, solution: true -->
YOUR ANSWER HERE

**7 (T&B 18.2). [10%]** Social scientists depend on the technique of _regression_, in which a vector of observations of some quantity is approximated in the least squares sense by a linear combination of other vectors. The coefficients of the fit are then interpreted as representing, say, the effects on annual income of IQ, years of education, parents' years of education, and parents' income.

One might think that the more variables one included in such model, the more information one would obtain, but this is not always true.  Explain this phenomenon from the point of view of conditioning, making specific reference to the results of T&B Theorem 18.1.

> **T&B Theorem 18.1.** Let $b \in \mathbb{C}^m$ and $A \in \mathbb{C}^{m \times n}$ of full rank be fixed. The least squares problem $\min_x \|b - Ax\|$ has the following 2-norm relative condition numbers describing the sensitivities of $y$ and $x$ to perturbations in $b$ and $A$:
$$\begin{array}{c|c|c|}
 & y & x \\ \hline
b & \frac{1}{\cos \theta} & \frac{\kappa(A)}{\eta \cos \theta} \\ \hline
A & \frac{\kappa(A)}{\cos \theta} & \kappa(A) + \frac{\kappa(A)^2\tan\theta}{\eta} \\ \hline
\end{array}.
$$
Here $\kappa(A) = \|A\|\|A^+\|$, $\theta = \cos^{-1} \frac{\|y\|}{\|b\|}$, $y$ is the projection of $b$
onto the range of $A$, and $\eta = \|A\|\,\|x\| / \|y\|.$
The results in the first row are exact, being attained for certain perturbations $\delta b$, and the results in the second row are upper bounds.

<!-- grade_id: cell-73d2c07a1958ed3c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

**8 (Based on Example from T&B Lecture 19).** Lecture 19 says that using modified Gram-Schmidt to solve the least squares problem _directly_ -- forming $\hat{Q}$ and $\hat{R}$ and computing $x \gets R \backslash Q^* b$ --
is unstable, but solving it _implicitly_ -- forming a reduced QR factorization $\tilde{Q}, \tilde{R}$ of the _augmented_ matrix
$[A | b]$, and computing $x \gets \tilde{R}_{1:n,1:n} \backslash \tilde{Q}_{1:n,n+1}$ --- is backward stable.

In this problem you will demonstrate this difference.

An implementation of modified Gram-Schmidt to compute a reduced QR factorization is provided.

```julia
function mgs(A)
    """
    Compute and return the QR factorization of `A` by Modified Gram-Schmidt procedure.

    # Arguments
    
    - `A`: m × n matrix (m ≥ n) to factor

    # Returns

    - `Q`: m × n matrix with orthonormal columns spanning range(A) 
    - `R`: n × n upper triangular matrix, coefficients of columns of `A` 
           in the basis of `Q`
    """
    _, n = size(A)
    Q = copy(A)
    R = zeros(eltype(A), n, n)
    for i in 1:n
        R[i,i] = norm(Q[:,i])
        Q[:,i] /= R[i,i]
        R[i,i+1:n] = Q[:,i]' * Q[:,i+1:n]
        Q[:,i+1:n] -= Q[:,[i]] * R[[i],i+1:n]
    end
    return Q, R
end
```

The function `generate_ls_problem` creates a least squares problem instance for testing.  It takes as inputs the size of $A$, as well as the factors that control the relative errors of $x$ and $y = Pb$:

- The condition number $\kappa(A)$
- The angle $\theta$ of $b$ to $\mathrm{range}(A)$ 
- The ratio $\eta$ of $\|x\|$ to $\|y\|$

```julia
function generate_ls_problem(m, n, κ, θ, η)
    """
    Create a random least squares problem with the desired conditioning factors
    
    # Arguments
    
    - `m`, `n`: the dimensions of the least squares problem
    - `κ`: the 2-norm condition number of `A`
    - `θ`: the angle of `b` to range(A)
    - `η`: the ratio of ||x|| to ||y||
    
    # Returns
    
    - `A`: the system matrix
    - `x`: the solution vector
    - `r`: the residual vector
    
    Compute `y = A*x` and `b = y + r` to construct the right hand side of the least squares problem
    """
    @assert(m >= n)
    @assert(κ >= 1.0)
    @assert(η >= 1.0 && κ >= η)
    
    # generate a random square matrix with the desired
    # condition number
    σ = LinRange(1.0, 1.0 / κ, n)
    U = qr(randn(n,n)).Q
    V = qr(randn(n,n)).Q
    A = zeros(m, n)
    A[1:n,1:n] = U * diagm(σ) * V'
    
    # create a solution vector with norm η cos θ
    x = zeros(n)
    ϕ = 1 / η
    # spread the support of x out widely over the singular vectors
    # pairing large and small singular values appropriately scaled
    # to guarantee ||x|| = η ||y||
    for i = 1:div(n,2)
        s = σ[i]
        t = σ[end-(i-1)]
        if (s >= ϕ && t <= ϕ)
            if s == t
                x[i] = 1.
                x[end-(i-1)] = 1.
            else
                c = (ϕ+t)*(ϕ-t)/(s+t)/(s-t)
                x[i] = sqrt(c)
                x[end-(i-1)] = sqrt(1 - c)
            end
        end
    end
    x = V * x
    # normalize x so that ||y|| = 1
    xnorm = norm(x)
    x *= η / xnorm
    
    # set up a right hand side where ||b|| = 1 and ||y = Pb|| = cos θ
    c = cos(θ)
    s = sin(θ)
    # rescale x so ||y|| = cos θ
    x *= c
    r = zeros(m)
    # put sin θ in the last component (this will be the residual r)
    r[end] = s
    
    # randomly rotate the columns
    Q = qr(randn(m,n)).Q
    A = Q'*A
    r = Q'*r
    
    return A, x, r
end
```

**(a) [6%]** Implement a function that solves the least squares problem directly as described above.

<!-- grade_id: cell-ac2ffa717e6d0b08, points: 0, grade: false, solution: true -->
```julia
function solve_ls_mgs_direct(A, b)
    """
    Solve minₓ ||Ax - b||₂ by direct application of modified Gram-Schmidt
    
    # Arguments
    
    - `A`: system matrix
    - `b`: right-hand side
    
    # Returns
    
    - `x`: the computed solution vector
    - `y`: the projection `Pb` of `b` onto the range of `A`
    """
    # your code here
    return x, y
end
```

<!-- grade_id: cell-0f942c43598294ea, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_direct on a random well-conditioned problem

m = 100; n = 15; κ = 10; θ = π / 20; η = κ / 2;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_direct(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by direct MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err < 100 * eps())
@assert(y_rel_err < 100 * eps())
```

<!-- grade_id: cell-fe782e442e467e73, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_direct on a moderately-conditioned problem with η = κ: stability is not a concern yet

m = 100; n = 15; κ = 1.e6; θ = π / 20; η = κ;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_direct(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by direct MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err < 100 * κ * eps())
@assert(y_rel_err < 100 * κ * eps())
```

<!-- grade_id: cell-474de05fca424d20, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_direct on an ill-conditioned problem with η = 2, and a small θ: should be unstable

m = 100; n = 15; κ = 1.e10; θ = 1.e-6; η = 2;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_direct(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by direct MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err > 1.0)
@assert(y_rel_err < 100 * κ * eps())
```

**(b) [6%]** Now implement a function that solve the least squares problem implicitly as described above.

<!-- grade_id: cell-fa1a67a405598d48, points: 0, grade: false, solution: true -->
```julia
function solve_ls_mgs_implicit(A, b)
    """
    Solve minₓ ||Ax - b||₂ by implicit application of modified Gram-Schmidt to the augmented system [A|b]
    
    # Arguments
    
    - `A`: system matrix
    - `b`: right-hand side
    
    # Returns
    
    - `x`: the computed solution vector
    - `y`: the projection `Pb` of `b` onto the range of `A`
    """
    # your code here
    return x, y
end
```

<!-- grade_id: cell-841ce780372c0568, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_implicit on a random well-conditioned problem

m = 100; n = 15; κ = 10; θ = π / 20; η = κ / 2;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_implicit(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by implicit MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err < 100 * eps())
@assert(y_rel_err < 100 * eps())
```

<!-- grade_id: cell-9739ac0d690e67ae, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_implicit on a moderately-conditioned problem with η = κ: still stable

m = 100; n = 15; κ = 1.e6; θ = π / 20; η = κ;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_implicit(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by implicit MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err < 100 * κ * eps())
@assert(y_rel_err < 100 * κ * eps())
```

<!-- grade_id: cell-0364b84a8fe406df, points: 2, grade: true, solution: false -->
```julia
# Test solve_ls_mgs_implicit on an ill-conditioned problem with η = 2, and a small θ: still stable
m = 100; n = 15; κ = 1.e10; θ = 1.e-6; η = 2;
A, x, r = generate_ls_problem(m, n, κ, θ, η)
y = A * x
b = r + y
x_dir, y_dir = solve_ls_mgs_implicit(A, b)
x_rel_err = norm(x_dir - x) / norm(x)
y_rel_err = norm(y_dir - y) / norm(y)

display("Solution by implicit MGS: (x, y) relative errors: ($x_rel_err, $y_rel_err)")
@assert(x_rel_err < 100 * m * n * κ * eps())
@assert(y_rel_err < 100 * κ * eps())
```

**(c) [4%]** Create two plots, one for the relative $x$ error $\|\hat x - x\| / \|x\|$ and one for the relative $y$ error $\|\hat y - y \| / \|y\|$.  Generate 10 least squares problems for each of the condition numbers $\kappa(A) = 10^1, 10^2, \dots, 10^10$, with $\theta = 10^{-6}$ and $\eta = \sqrt{\kappa}$.
Solve each generated least squares problem directly and implicitly and record the $x$ and $y$ relative errors.
Plot the relative errors in the $x$ and $y$ plots as scatter plots.  Use $\kappa$ as the x-axis and the relative error as the y-axis.  Use logarithmic scales for both axes.  Make sure the difference between the direct and implicit data points is clearly labeled in your plots.

<!-- grade_id: cell-fe16ee496bcfa924, points: 4, grade: true, solution: true -->
```julia
# your code here
```