# Homework 4: Systems of Equations

```julia
import Pkg; Pkg.add("Plots");
using LinearAlgebra; using Plots; default(fmt=:png);
BLAS.set_num_threads(1);
```

```julia
bigspy(A; kwargs...) = spy(A, m=4; kwargs...)
```

## 1. (T&B Exercise 20.2)

Suppose $A \in \mathbb{C}^{m \times m}$ has an LU factorization (by T&B Exercise 20.1, this is true if and only if for each $1 \leq k \leq m$ the upper-left $k \times k$ block $A_{1:k,1:k}$ is nonsingular) and is banded with bandwidth $2p + 1$, i.e., $a_{ij} = 0$ for $|i - j| > p$.

### 1.(a) [5%]

Create a random $20 \times 20$ matrix $A$ with bandwidth $2p + 1 = 5$, i.e. $p=2$, and compute its LU factorization $A = LU$ (without pivoting: see the [documentation](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.lu) to determine how to turn off pivoting in `lu`).  Plot the matrices $A$, $L$ and $U$ using the `bigspy` command.

<!-- grade_id: cell-7a94a0178a69658f, points: 5, grade: true, solution: true -->
```julia
# your code here
```

### 1.(b) [5%]

Without adding the plots here, try the same for different values of $p$ until you see a pattern in the sparsity patterns of $L$ and $U$.  State what that pattern is and prove that it holds for general $p$.

<!-- grade_id: cell-09e2910abafc0e71, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## 2. (T&B Exercise 21.2)

### 2.(a) [5%]

Repeat 1.(a), but this time with partial pivoting, and generate the same three plots.

<!-- grade_id: cell-03674065c637d279, points: 5, grade: true, solution: true -->
```julia
# your code here
```

### 2.(b) [5%]

What can be said about the bandedness of $U$?  Specifically, you should determine $p_U$ such that the bandwidth of $U$ is at most $2p_U + 1$ as a function of $p$. (Again, you may want to test your result on several random inputs until you see a pattern.) Prove your result.

<!-- grade_id: cell-a316459aa5562639, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 2.(c) [5%]

What is the most nonzeros entries $n_L$ that are in any column of $L$? Prove your result.

<!-- grade_id: cell-03e1a326ee6a5a6c, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## 3. (T&B Exercise 20.4) [10%]

Write the LU factorization algorithm (without pivoting) using just one explicit `for` loop.  Inside this loop, $U$ will be updated at each step by a certain rank-one outer product.  This "outer product" form of Gaussian elimination may be a better starting point if one wants to optimize computer performance.

<!-- grade_id: cell-464a56320e48763f, points: 0, grade: false, solution: true -->
```julia
function lu_one_loop(A)
    """
    Compute the LU factorization of A using one explicit for loop.
    
    Arguments:
    
    - `A`: A square matrix
    
    Returns:
    
    - `L`: A unit lower triangular matrix
    
    - `U`: An upper triangular matrix such that A = L*U
    """
    # your code here
    return L, U
end
```

```julia
function generate_random_lu_data(T, m)
    """
    Generate a random matrix that is guaranteed to have an LU factorization without pivoting.
    
    Arguments:
    
    - `T`: A datatype
    - `m`: A positive integer
    
    Returns:
    
    - `A`: an m × m matrix that has an LU factorization
    """
    
    A = randn(T, m, m)
    col_sum = abs.(A)' * ones(m)
    signs = sign.(diag(A))
    A += 1.1 * diagm(col_sum .* signs)
    return A
end
```

```julia
# Test on real inputs
m = 15
A = generate_random_lu_data(Float64, m);
```

<!-- grade_id: cell-87f388177392b442, points: 1, grade: true, solution: false -->
```julia
# Q3: Run `lu_one_loop` on real inputs
L, U = lu_one_loop(A);
```

<!-- grade_id: cell-7058378c4d71e098, points: 1, grade: true, solution: false -->
```julia
# Q3: Real L is lower triangular
@assert(norm(L - LowerTriangular(L)) == 0)
```

<!-- grade_id: cell-6e896b73af0cbe8d, points: 1, grade: true, solution: false -->
```julia
# Q3: Real L is unit lower triangular
@assert(norm(diag(L) - ones(m)) == 0)
```

<!-- grade_id: cell-44493dd86a89244b, points: 1, grade: true, solution: false -->
```julia
# Q3: Real U is upper triangular
@assert(norm(U - UpperTriangular(U)) == 0)
```

<!-- grade_id: cell-a399ecee00906e90, points: 1, grade: true, solution: false -->
```julia
# Q3: Real LU is close to A
ρ_est = sqrt(m) # average case observed growth factor
@assert(norm(A - L*U) <= 100 * eps() * ρ_est * norm(A))
```

```julia
# Test on complex inputs
m = 15
A = generate_random_lu_data(Complex{Float64}, m);
```

<!-- grade_id: cell-2ed27e7af70a2b70, points: 1, grade: true, solution: false -->
```julia
# Q3: Run `lu_one_loop` on complex inputs
L, U = lu_one_loop(A);
```

<!-- grade_id: cell-c7dc0c784ab84aec, points: 1, grade: true, solution: false -->
```julia
# Q3: Complex L is lower triangular
@assert(norm(L - LowerTriangular(L)) == 0)
```

<!-- grade_id: cell-c2802feabed8c255, points: 1, grade: true, solution: false -->
```julia
# Q3: Complex L is unit lower triangular
@assert(norm(diag(L) - ones(m)) == 0)
```

<!-- grade_id: cell-8eedf3c72fd4fd1f, points: 1, grade: true, solution: false -->
```julia
# Q3: Complex U is upper triangular
@assert(norm(U - UpperTriangular(U)) == 0)
```

<!-- grade_id: cell-1675a69b08de2c4b, points: 1, grade: true, solution: false -->
```julia
# Q3: Complex LU is close to A
ρ_est = sqrt(m) # average case observed growth factor
@assert(norm(A - L*U) <= 100 * eps() * ρ_est * norm(A))
```

## 4. (T&B Exercise 21.6) [10%]

You'll notice that `generate_random_lu_data` first creates a matrix and then adds a multiple of the column sum
$\sum_{i} |a_{ij}|$ to the diagonal $a_{jj}$, and that the result should have an LU factorization without pivoting.

Suppose $A \in \mathbb{C}^{m \times m}$ is _strictly column diagonally dominant_, which means that for each $k$,

$$|a_{kk}| > \sum_{j \neq k} |a_{jk}|.$$

Show that if Gaussian elimination with partial pivoting is applied to $A$, no row interchanges take place.

<!-- grade_id: cell-ff12f6c1891b4f17, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## 5. (T&B Exercise 22.2)

### 5.(a) [5%]

Write a function that for a given $m$ generate an $m \times m$ matrix $A$ for which the backward error of
pivoted LU factorization, $\tilde{L}\tilde{U} = \tilde{P} A + \delta A$ grows like the pessimistic bound

$$\frac{\max_{ij} |(\delta A)_{ij}|}{\max_{ij} |A_{ij}|} = O(2^{m-1} \epsilon_{\text{machine}}).$$

In particular, your method should generate matrices such that the following bounds hold:

| $m$ | $\max_{ij} |(\delta A)_{ij}| / \max_{ij} |A_{ij}|$ |
| :-: | :-: |
| $30$  | $\geq 2^{20} \epsilon_{\text{machine}}$ |
| $40$  | $\geq 2^{30} \epsilon_{\text{machine}}$ |
| $50$  | $\geq 2^{40} \epsilon_{\text{machine}}$ |

<!-- grade_id: cell-c5e3dd2fb51ec8a2, points: 2, grade: true, solution: true -->
```julia
function catastrophic_growth_factor_matrix(m)
    """
    Construct a matrix with nearly pessimistic growth factor
    
    Arguments:
    
    - `m`: An integer
    
    Returns:
    
    - `A`: An m × m matrix whose growth factor (computed with julia's `lu`) is close
           to the upper bound 2ᵐ⁻¹.
    """
    # your code here
    return A
end
```

<!-- grade_id: cell-281155809b25fb3d, points: 1, grade: true, solution: false -->
```julia
# Q5: test m = 30
m = 30
A = catastrophic_growth_factor_matrix(m)
f = lu(A)
rel_err = norm(f.P * A - f.L * f.U, Inf) / norm(A, Inf)
@assert rel_err >= eps() * 2.0 ^ (m - 10)
```

<!-- grade_id: cell-ca18063920244d94, points: 1, grade: true, solution: false -->
```julia
# Q5: test m = 40
m = 40
A = catastrophic_growth_factor_matrix(m)
f = lu(A)
rel_err = norm(f.P * A - f.L * f.U, Inf) / norm(A, Inf)
@assert rel_err >= eps() * 2.0 ^ (m - 10)
```

<!-- grade_id: cell-1d4b2929db04fb3e, points: 1, grade: true, solution: false -->
```julia
# Q5: test m = 50
m = 50
A = catastrophic_growth_factor_matrix(m)
f = lu(A)
rel_err = norm(f.P * A - f.L * f.U, Inf) / norm(A, Inf)
@assert rel_err >= eps() * 2.0 ^ (m - 10)
```

### 5.(b) [5%]

Make a plot of the backward relative errors $\max_{ij} |(\delta A)_{ij}| / \max_{ij} |A_{ij}|$ for your matrices for values $m=10, 11, \dots, 60$.  Make $m$ the x-axis and the relative error the y-axs, and make the y-axis a logarithmic scale.
Add a trend line for $2^{m-1} \epsilon_{\text{machine}}$ to your plot to show what the asymptotic growth looks like in the worst case.

<!-- grade_id: cell-84f869dfe19de685, points: 5, grade: true, solution: true -->
```julia
# your code here
```

## 6. (T&B Exercise 22.4)

### 6.(a) [5%]

Suppose $PA = LU$ (LU factorization with partial pivoting) and $A = QR$ (QR factorization).  Describe a relationship between the last row of $L^{-1}$ and the last column of $Q$, and prove why this relationship is so.

<!-- grade_id: cell-9ae93b05202f8fd7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 6.(b) [5%]

Show that if $A$ is random in the sense of having independent, normally distributed entries, then its column spaces are randomly oriented, so that in particular, the last column of $Q$ is a random unit vector.

<!-- grade_id: cell-b8219f8a48b99bdf, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 6.(c) [5%]

Combine the results of (a) and (b) to make a statement about the final row of $L^{-1}$ in Gaussian elimination applied to a random matrix $A$.

<!-- grade_id: cell-81c1dbc190d6dd86, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## 7.

### 7.(a) [5%]

Give an example of a $3 \times 3$ matrix $A$ in block form

$$
A = \begin{bmatrix} F & B \\ B^* & 0 \end{bmatrix},
$$

where:

- $F$ is a $2 \times 2$ symmetric positive definite matrix and $\kappa(F) \geq 10^{10}$, and
- $\kappa(A) < 2$.

<!-- grade_id: cell-79dd2b537400d961, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## 7.(b) [5%]

Consider the problem of solving the system of equations

$$
\begin{bmatrix}
C & B \\ B^* & 0
\end{bmatrix}
\begin{bmatrix}
x \\ y
\end{bmatrix}
=
\begin{bmatrix}
f \\ g
\end{bmatrix},
$$

where $C$ is symmetric positive definite and $B \in \mathbb{C}^{m \times n}$ ($m \geq n$) has full rank.  Suppose we solve this system in the following steps:

1. Compute the Cholesky factorization $C = LL^*$.
2. Compute $w = g - B^*C^{-1} f = g - B^* L^{-*} L^{-1} f$.
3. Compute the matrix $S= B^* C^{-1} B = B^* L^{-*} L^{-1} B$.
4. Compute the Cholesky factorization $G G^* = S$.
5. Compute $y = -S^{-1} w = - G^{-*} G^{-1} w$.
6. Compute $z = f - B y$.
7. Compute $x = C^{-1} z = L^{-*} L^{-1}z$.

Why does part $A$ suggest that this is method may be unstable?

<!-- grade_id: cell-d8bb8ec0c2956260, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 7.(c) [10%]

Now suppose we consider only problems where $g = 0$,

$$
\begin{bmatrix}
C & B \\ B^* & 0
\end{bmatrix}
\begin{bmatrix}
x \\ y
\end{bmatrix}
=
\begin{bmatrix}
f \\ 0
\end{bmatrix},
$$

so that $y = (B^* C^{-1} B)^{-1} B^* C^{-1} f.$

Prove that there is a bound $\psi_B$ such that

$$\|(B^* C^{-1} B)^{-1} B^* C^{-1}\| \leq \psi_B, \quad \text{for all symmetric positive definite }C.$$

This suggests that the condition number of $\kappa_{C \to y}$ is independent of $C$.

_Hint:_ Use two SVDs: one of $B$, one of $B^* C^{-1}$.

<!-- grade_id: cell-079c2be0ae362ecf, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## 8. 

### 8.(a) [6%]

Write an implementation of a rank-one Cholesky update: assuming $L L^* = A$ is a Cholesky factorization of $A$,
compute $H H^* = A + u u^*$, the Cholesky factorization of the rank-one update of $A$.

Unlike a Cholesky factorization from scratch, your implementation should use $O(m^2)$ operations.

You are welcome to consult [existing implementations](https://en.wikipedia.org/wiki/Cholesky_decomposition#Updating_the_decomposition), but your implementation will have to be valid for complex-valued $A$ and $u$.

<!-- grade_id: cell-b47ad90c4e69b139, points: 0, grade: false, solution: true -->
```julia
function cholesky_update(L, u)
    """
    Update a Cholesky factorization by a rank one update,
    
        H H' = L L' + u u'
    
    Arguments:
    
    - `L`: An existing lower-triangular Cholesky factor
    
    - `u`: The vector for the rank-one update
    
    Returns:
    
    - `H`: the lower-triangular Cholesky factor of `L L' + u u'`
    
    """
    # your code here
    return H
end
```

<!-- grade_id: cell-748716e50e40e09e, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Run cholesky_update on real input
m = 50
B = randn(m,m)
A = B' * B
L = cholesky(A).L
u = randn(m)
Au = A + u*u'
H = cholesky_update(L, u);
```

<!-- grade_id: cell-6fad27a8867dd8a9, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Real Cholesky factor is lower triangular
@assert(norm(H - LowerTriangular(H)) == 0)
```

<!-- grade_id: cell-44ed8b724f8f7116, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Real Cholesky factor has small backward error
@assert(norm(Au - H * H') / norm(Au) <= eps() * 100)
```

<!-- grade_id: cell-4add80ef39d6201b, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Run cholesky_update on complex input
m = 50
B = randn(Complex{Float64}, m,m)
A = B' * B
L = cholesky(A).L
u = randn(Complex{Float64}, m)
Au = A + u*u'
H = cholesky_update(L, u);
```

<!-- grade_id: cell-84d61897527c1d54, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Complex Cholesky factor is lower triangular
@assert(norm(H - LowerTriangular(H)) == 0)
```

<!-- grade_id: cell-f5d9fe490aaee5e3, points: 1, grade: true, solution: false -->
```julia
# Q8.(a): Complex Cholesky factor has small backward error
@assert(norm(Au - H * H') / norm(Au) <= eps() * 100)
```

### 8.(b) [4%]

Suppose that we only want to solve the system $(A + u u^*) x = b$.

Again assume that $L L^* = A$ is a Cholesky factorization of $A$, but instead of updating the factorization, write a function that solves the system using the [Woodbury formula](https://en.wikipedia.org/wiki/Woodbury_matrix_identity).

<!-- grade_id: cell-f07c1f58376e3ce3, points: 0, grade: false, solution: true -->
```julia
function solve_cholesky_update(L, u, b)
    """
    Solve (L L' + u u') x = b
    
    Arguments:
    
    - `L`: An existing lower-triangular Cholesky factor
    
    - `u`: The vector for the rank-one update
    
    Returns:
    
    - `x` : the solution of (L L' + u u') x = b
    """
    # your code here
    return x
end
```

<!-- grade_id: cell-b4cdd9f4a74f9de3, points: 1, grade: true, solution: false -->
```julia
# Q8.(b): Run solve_cholesky_update on real input
m = 50
B = randn(m,m)
A = B' * B
L = cholesky(A).L
u = randn(m)
b = randn(m)
Au = A + u*u'
x = solve_cholesky_update(L, u, b);
```

<!-- grade_id: cell-acd58ca7852fe296, points: 1, grade: true, solution: false -->
```julia
# Q8.(b): Real cholesky update solve has small backward error
@assert(norm(Au * x - b) / norm(b) <= eps() * 1.e5)
```

<!-- grade_id: cell-a8c9cda055f65b4b, points: 1, grade: true, solution: false -->
```julia
# Q8.(b): Run solve_cholesky_update on complex input
m = 50
B = randn(Complex{Float64}, m,m)
A = B' * B
L = cholesky(A).L
u = randn(Complex{Float64}, m)
b = randn(Complex{Float64}, m)
Au = A + u*u'
x = solve_cholesky_update(L, u, b);
```

<!-- grade_id: cell-87045cc68fb4b6cc, points: 1, grade: true, solution: false -->
```julia
# Q8.(b): Complex cholesky update solve has small backward error
@assert(norm(Au * x - b) / norm(b) <= eps() * 1.e5)
```